# EEN065 - Applied object-oriented programming


Teacher: [Seyed Morteza Ahmadian](https://www.chalmers.se/en/persons/seyedah/) and [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/39015)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before submitting this assignment, please ensure that everything runs correctly.

First, **restart the kernel**, then **run all cells**.
Next, verify that all tests execute successfully.

Note that if any problem produces an error, the subsequent problems **will not** be tested.

In case of discrepancies between the problem description and the tests, you should resolve the issue with the tests in mind, as they determine correctness.

There are four types of cells in this notebook:
1.	*Solution cells*: These are the cells where you write your answers or modify the existing code to solve the problem.
2.	*Debug cells*: These cells are for debugging your solution. You may add or edit code here, but ensure that it does not introduce any errors.
3. *Test cells:*  These cells run **automated tests** to verify whether your solution is **functionally correct**.
4. *Code quality tests:*  These cells check whether your solution follows the **code quality standards** of the course. Even if your code produces the correct output, it must also be **clean, readable, and well-structured**.

- If all tests pass, you will see an output similar to:
  - `1 passed`

- If a test fails, you will see an **error message** explaining:
  - Which input was tested
  - What output was expected
  - What your function actually returned

You will see several sections in the output:

- **tests:** Runs the test cells with coverage enabled.  
  - `1 passed` indicates success.
- **coverage:** Checks for any line of your solution that is not useful.  
  - `100% coverage` indicates success.
- **flake8:** Checks for style issues (e.g., unused variables, line length).  
  - Any output other than `0` here means you must fix style violations.
- **pylint:** Performs deeper code analysis and assigns a score.  
  - A score of `10.00/10` means your code fully complies.
- **mypy**: Checks that your code respects the specified type annotations.
  - Any reported error means a type hint is missing or incorrect.
- **isort:** Verifies that imports are correctly ordered.  
  - If differences are shown, reorder your imports.

At the end:
- **`tests passed ✓`** → Your solution is correct *and* meets all quality requirements.
- **`your code presented some issue`** → Review the messages above, fix the reported issues, and rerun the cell.

**You must pass both functional tests and code quality checks to receive full credit.**

Common errors and what they mean:
- **`You did not execute your solution cell!`**  
  → You forgot to run the solution cell before running the test cell.
- **`The module does not have the necessary function!`**  
  → The required function is missing or named incorrectly.
- **AssertionError**  
  → Your function ran, but it returned an incorrect result for at least one test case.

**Fix all test errors before moving on**, as later problems will not be tested if an earlier one fails.

**Delete** the line `raise NotImplementedError()` from any problem you solve.

**Do not add or remove any cells in this file.** All required cells are already provided.

To execute a cell, select it and press **CTRL+Enter** (Windows), **CMD+Enter** (macOS), or click the **Run cell** button.

## Assignment Week 3

**Observations:**
For the tasks in this assignment, you should:
- use only the math operations studied in the lecture notes (i.e., no advanced functions allowed).
- not use external modules, i.e., `import` command is not allowed in your solutions (with exception of the module `typing`).

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
import os

# getting latest version of utils.py
# import requests
# response = requests.get("https://raw.githubusercontent.com/carlosnatalino/programming-assignments/refs/heads/main/utils.py")
# with open("utils.py", "wb") as f:
#     f.write(response.content)

if not os.path.exists("utils.py"):
    raise ValueError("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.")

# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import export_cell_content
except:
    import sys
    raise ValueError("Did you select the right kernel that starts with venv?")

filename = "assignment_week_03_public.ipynb"
assert filename != None, "The file name must be set."
if not os.path.exists(filename):
    raise ValueError(f"The value of the `filename` variable is wrong. \
The file `{filename}` does not exist. Write the correct name of this file!")

os.environ["NOTEBOOK_NAME"] = os.path.join(os.getcwd(), filename)
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

### Task 1

Create a function called `assignment_points` that receives an integer representing the number of assignment approved.
The function then returns a float that represents the assignment points obtained according to the scale defined in the [course PM](https://chalmers.instructure.com/courses/39015#:~:text=1%20AP,15%20APs).

Make sure to return `0` if no assignment is approved, i.e., when you recieve the input `0`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def assignment_points(n: int) -> float:
    points = [0.0, 1.0, 2.0, 4.0, 6.0, 9.0, 12.0, 15.0]
    return points[n]
```

```python
# uncomment the line(s) below to debug

count = 6
ap = assignment_points(count)
print(f"The assignment points are {ap}")
print("Execution finished", "\u2713")
```

Output:
```
The assignment points are 12.0
Execution finished ✓
```

```python
# test cell
def test_solution_1() -> None:
    try:
        import solution_1
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_1 import assignment_points
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(6, 12.0), (2, 2.0), \
                (7, 15.0), (3, 4.0), \
                (4, 6.0),  (5, 9.0), \
                (1, 1.0), (0, 0.0)]

    for _in, _out in test_cases:
        _res = assignment_points(_in)
        assert _res == _out, f"The function with input `{_in}` should return the value \
    `{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_1.py"
test_file = "test_solution_1.py"
export_cell_content("solution_1", solution_file)
export_cell_content("tests_1", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:01:23.619158 (2 seconds ago)
Latest saved date of the file: 2026-04-16 14:01:23.619158 (2 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.05s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_1.py       3      0      0      0   100%
-----------------------------------------------------------
TOTAL               3      0      0      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ################
```

### Task 2

Create a function called `smallest_value` that receives three numbers (`int` or `float`) as an input, and returns the smallest one among them.

**Note:**
- You must not use built-in functions to solve the problem; use only conditional statements.
- You must not use any import statement.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
# solution cell
def smallest_value(a: float, b: float, c: float) -> float:
    if a <= b and a <= c:
        return a
    if b <= a and b <= c:
        return b
    return c
```

```python
# uncomment the line(s) below to debug

a, b, c = 12, 12, 11
l = smallest_value(a, b, c)
print(f"The smallest value among {a}, {b}, {c} is {l}")
print("Execution finished", "\u2713")
```

Output:
```
The smallest value among 12, 12, 11 is 11
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_2() -> None:
    try:
        import solution_2
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_2 import smallest_value
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((1, 2.2, 3), 1), ((2, 1.1, 3), 1.1), ((3, 2.4, 2.3), 2.3), \
    ((1, -2, -7), -7)]

    for _in, _out in test_cases:
        _res = smallest_value(*_in)
    assert _res == _out, f"The function with input `{_in}` should return the value \
`{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_2.py"
test_file = "test_solution_2.py"
export_cell_content("solution_2", solution_file)
export_cell_content("tests_2", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:01:23.619158 (5 seconds ago)
Latest saved date of the file: 2026-04-16 14:01:23.619158 (5 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.06s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_2.py       6      0      4      0   100%
-----------------------------------------------------------
TOTAL               6      0      4      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ################
```

### Task 3

Create a Python function named `reverse_lower_and_upper_case` that receives a string as an input. It reverses the case of the input. That is if the input is in upper case, then convert it to lower case and if it is in lower case, then convert it to lower case.

**Examples**

Input: 'test'

Output: 'TEST'

Explanation: Here 'test' is in lower case, so we converted it to upper case.

**Hint:** You must check the `string` official Python documentation linked at the end of Chapter 5.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
# solution cell
def reverse_lower_and_upper_case(x: str) -> str:
    if x.islower() is True:
        return x.upper()
    return x.lower()
```

```python
# uncomment the line(s) below to debug

initial = ['test', 'PRACTICE']
for _i in initial:
    result = reverse_lower_and_upper_case(_i)
    print(f"The reverse case version of {_i} is {result}")
print("Execution finished", "\u2713")
```

Output:
```
The reverse case version of test is TEST
The reverse case version of PRACTICE is practice
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_3() -> None:
    try:
        import solution_3
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_3 import reverse_lower_and_upper_case
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [('test', 'TEST'), \
        ('practice', 'PRACTICE'), \
        ('PRACTICE', 'practice'), \
        ('RESOURCE NOT GRANTED', 'resource not granted')]

    for _in, _out in test_cases:
        _res = reverse_lower_and_upper_case(_in)
        assert _res == _out, f"The function should return the value `{_out}` of type \
    `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
    print("tests passed", "\u2713")
```

```python
# validate the solution
solution_file = "solution_3.py"
test_file = "test_solution_3.py"
export_cell_content("solution_3", solution_file)
export_cell_content("tests_3", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:01:23.619158 (7 seconds ago)
Latest saved date of the file: 2026-04-16 14:01:23.619158 (7 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.07s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_3.py       4      0      2      0   100%
-----------------------------------------------------------
TOTAL               4      0      2      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ################
```

### Task 4

Create a Python function called `mobile_number_validate` that receives a mobile number as a string and validates whether it is a correct mobile number (returning `True`) or not (returning `False`). 

Rules for mobile number:
- string starts with either of `0`, `46`, or `+46`;
- the remaining string has length 9 such that the first digit is between 1 to 9 and the other digits are between 0 to 9.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-4, points: 0, grade: false, solution: true -->
```python
# solution cell
def mobile_number_validate(number: str) -> bool:
    if number[0] == "0":
        rest = number[1:]
    elif number[0:2] == "46":
        rest = number[2:]
    elif number[0:3] == "+46":
        rest = number[3:]
    else:
        return False
    if len(rest) != 9:
        return False
    if rest[0] == "0":
        return False
    numb = 0
    while numb < len(rest):
        if rest[numb] < "0" or rest[numb] > "9":
            return False
        numb += 1
    return True
```

```python
# uncomment the line(s) below to debug

number = "46102345678" # Valid
valid = mobile_number_validate(number)
valid_string = "valid" if valid else "invalid"
print(f"The mobile number {number} is {valid_string}")
print("Execution finished", "\u2713")
```

Output:
```
The mobile number 46102345678 is valid
Execution finished ✓
```

<!-- grade_id: cell-test-1-4, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_4() -> None:
    try:
        import solution_4
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_4 import mobile_number_validate
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [('0061234567', False), ('+46061234567', False), ('123456789', False), \
    ('+465555', False), ('46123456789', True), ('+46761234567', True)]

    for _in, _out in test_cases:
        _res = mobile_number_validate(_in)
        assert _res == _out, f"The function with input `{_in}` should return the value \
    `{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_4.py"
test_file = "test_solution_4.py"
export_cell_content("solution_4", solution_file)
export_cell_content("tests_4", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

# print('\n\033[0m')
# print(" coverage ".center(100, "#"))
# !coverage report -m {solution_file}
# if _exit_code != 0:
#     passed = False
# print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:01:23.619158 (10 seconds ago)
Latest saved date of the file: 2026-04-16 14:01:23.619158 (10 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.07s[0m[0m

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ##############################################

--------------------------------------------------------------------
Your code has been rated at 10.00/10 (previous run: 10.00/10, +0.00)

pylint True

[0m
############################################### mypy ###############################################
[1m[92mSuccess: no issues found in 1 source file[0m
mypy True

[0m
######################################
```

### Task 5

Create a Python function called `sum_multiples` that receives two parameters:

- `n: int`
- `k: int`

The function must compute the sum of all numbers from `1` to `n` (inclusive) that are divisible by `k`.

**Note:**
- If `n` is less than `1`, return `0`.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def sum_multiples(n: int, k: int) -> int:
    summa = 0
    for x in range(1, n + 1):
        if x % k == 0:
            summa += x
    return summa
```

```python
# uncomment the line(s) below to debug

print(sum_multiples(10, 3))  # expected 18
print(sum_multiples(5, 2))   # expected 6

print("Execution finished", u"\u2713")
```

Output:
```
18
6
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_5() -> None:
    try:
        import solution_5
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_5 import sum_multiples
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [
        ((10, 3), 18),
        ((5, 2), 6),
        ((1, 1), 1),
        ((0, 3), 0),
        ((-5, 2), 0),
        ((7, 10), 0),
    ]

    for _in, _out in test_cases:
        _res = sum_multiples(*_in)
        assert _res == _out, \
            f"The function with input `{_in}` should return `{_out}`, but returned `{_res}`."
```

```python
# validate the solution
solution_file = "solution_5.py"
test_file = "test_solution_5.py"
export_cell_content("solution_5", solution_file)
export_cell_content("tests_5", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:02:37.091424 (6 seconds ago)
Latest saved date of the file: 2026-04-16 14:02:37.091424 (6 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.09s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_5.py       6      0      4      0   100%
-----------------------------------------------------------
TOTAL               6      0      4      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ################
```

### Task 6

Create a Python function called `alternating_sum` that receives one parameter:

- `n: int`

The function must compute the sum of numbers from `1` to `n` (inclusive), where:

- **odd numbers are added**
- **even numbers are subtracted**

If `n` is less than `1`, return `0`.

#### Example

```python
alternating_sum(5) → returns 3
```

Explanation: ```1 - 2 + 3 - 4 + 5 = 3```

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def alternating_sum(n: int) -> int:
    summa = 0
    for x in range(1, n + 1):
        if x % 2 == 0:
            summa -= x
        else:
            summa += x
    return summa
```

```python
# uncomment the line(s) below to debug

print(alternating_sum(5))  # expected 3
print(alternating_sum(4))  # expected -2

print("Execution finished", u"\u2713")
```

Output:
```
3
-2
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_6() -> None:
    try:
        import solution_6
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_6 import alternating_sum
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [
        (5, 3),
        (4, -2),
        (1, 1),
        (2, -1),
        (0, 0),
        (-3, 0),
    ]

    for _in, _out in test_cases:
        _res = alternating_sum(_in)
        assert _res == _out, \
            f"The function with input `{_in}` should return `{_out}`, but returned `{_res}`."
```

```python
# validate the solution
solution_file = "solution_6.py"
test_file = "test_solution_6.py"
export_cell_content("solution_6", solution_file)
export_cell_content("tests_6", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-16 14:01:23.619158 (15 seconds ago)
Latest saved date of the file: 2026-04-16 14:01:23.619158 (15 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.06s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_6.py       7      0      4      0   100%
-----------------------------------------------------------
TOTAL               7      0      4      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ##############
```