# EEN065 - Applied object-oriented programming


Teacher: [Seyed Morteza Ahmadian](https://www.chalmers.se/en/persons/seyedah/) and [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/39015)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before submitting this assignment, please ensure that everything runs correctly.

First, **restart the kernel**, then **run all cells**.
Next, verify that all tests execute successfully.

Note that if any problem produces an error, the subsequent problems **will not** be tested.

In case of discrepancies between the problem description and the tests, you should resolve the issue with the tests in mind, as they determine correctness.

There are four types of cells in this notebook:
1.	*Solution cells*: These are the cells where you write your answers or modify the existing code to solve the problem.
2.	*Debug cells*: These cells are for debugging your solution. You may add or edit code here, but ensure that it does not introduce any errors.
3. *Test cells:*  These cells run **automated tests** to verify whether your solution is **functionally correct**.
4. *Code quality tests:*  These cells check whether your solution follows the **code quality standards** of the course. Even if your code produces the correct output, it must also be **clean, readable, and well-structured**.

- If all tests pass, you will see an output similar to:
  - `1 passed`

- If a test fails, you will see an **error message** explaining:
  - Which input was tested
  - What output was expected
  - What your function actually returned

You will see several sections in the output:

- **tests:** Runs the test cells with coverage enabled.  
  - `1 passed` indicates success.
- **coverage:** Checks for any line of your solution that is not useful.  
  - `100% coverage` indicates success.
- **flake8:** Checks for style issues (e.g., unused variables, line length).  
  - Any output other than `0` here means you must fix style violations.
- **pylint:** Performs deeper code analysis and assigns a score.  
  - A score of `10.00/10` means your code fully complies.
- **mypy**: Checks that your code respects the specified type annotations.
  - Any reported error means a type hint is missing or incorrect.
- **isort:** Verifies that imports are correctly ordered.  
  - If differences are shown, reorder your imports.

At the end:
- **`tests passed ✓`** → Your solution is correct *and* meets all quality requirements.
- **`your code presented some issue`** → Review the messages above, fix the reported issues, and rerun the cell.

**You must pass both functional tests and code quality checks to receive full credit.**

Common errors and what they mean:
- **`You did not execute your solution cell!`**  
  → You forgot to run the solution cell before running the test cell.
- **`The module does not have the necessary function!`**  
  → The required function is missing or named incorrectly.
- **AssertionError**  
  → Your function ran, but it returned an incorrect result for at least one test case.

**Fix all test errors before moving on**, as later problems will not be tested if an earlier one fails.

**Delete** the line `raise NotImplementedError()` from any problem you solve.

**Do not add or remove any cells in this file.** All required cells are already provided.

To execute a cell, select it and press **CTRL+Enter** (Windows), **CMD+Enter** (macOS), or click the **Run cell** button.

## Assignment Week 5

**Observations:**
For the problems in this assignment, you should:
- make sure that your `utils.py` file has the following line.
- you will find the set of references needed to solve this assignment in the assignment page in canvas.

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
import os

# getting latest version of utils.py
# import requests
# response = requests.get("https://raw.githubusercontent.com/carlosnatalino/programming-assignments/refs/heads/main/utils.py")
# with open("utils.py", "wb") as f:
#     f.write(response.content)

if not os.path.exists("utils.py"):
    raise ValueError("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.")

# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import export_cell_content
except:
    import sys
    raise ValueError("Did you select the right kernel that starts with venv?")

filename = "assignment_week_05_public.ipynb"
assert filename != None, "The file name must be set."
if not os.path.exists(filename):
    raise ValueError(f"The value of the `filename` variable is wrong. \
The file `{filename}` does not exist. Write the correct name of this file!")

os.environ["NOTEBOOK_NAME"] = os.path.join(os.getcwd(), filename)
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

### Task 1

Create a Python function called `group_by_first_letter` that receives a list of strings.

The function must return a dictionary where:
- Each key is the first letter of a word (case-sensitive).
- Each value is a list of words starting with that letter.
- Words must remain in their original order.
- Empty strings must be ignored.

Rules
- If the input list is empty, return an empty dictionary.
- You must use `defaultdict` from `collections` module. Do not manually implement key-initialization logic.
- You must import only the required functions. Do not import the entire module or use wildcard imports.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from collections import defaultdict


def group_by_first_letter(x: list[str]) -> dict[str, list[str]]:
    d = defaultdict(list)
    for word in x:
        if word == "":
            continue
        första = word[0]
        d[första].append(word)
    return dict(d)
```

```python
### uncomment the line(s) below to debug

input_list = ["apple", "banana", "apricot", "Berry", "", "blue"]
out = group_by_first_letter(input_list)
print(f"The dictionary grouped by first letter is\n{out}")
print("Execution finished", "\u2713")
```

Output:
```
The dictionary grouped by first letter is
{'a': ['apple', 'apricot'], 'b': ['banana', 'blue'], 'B': ['Berry']}
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch


def test_solution_1() -> None:
    try:
        import solution_1
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_1 import group_by_first_letter
    except:
        raise ValueError("The module does not have the necessary function!")

    words = ["apple", "banana", "apricot", "Berry", "", "blue"]

    result = group_by_first_letter(words)

    expected = {
        "a": ["apple", "apricot"],
        "b": ["banana", "blue"],
        "B": ["Berry"],
    }

    assert result == expected, \
        f"Expected {expected}, but got {result}."

    assert group_by_first_letter([]) == {}

    with patch("solution_1.defaultdict") as mock_defaultdict:
        mock_instance = mock_defaultdict.return_value
        mock_instance.__getitem__.side_effect = lambda key: []
        mock_instance.items.return_value = []

        group_by_first_letter(["x"])

        mock_defaultdict.assert_called_once()
        assert mock_defaultdict.called, \
            "Do not manually implement dictionary initialization logic. Use `defaultdict`."
```

```python
# validate the solution
solution_file = "solution_1.py"
test_file = "test_solution_1.py"
export_cell_content("solution_1", solution_file)
export_cell_content("tests_1", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 1.30s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_1.py       9      0      4      0   100%
-----------------------------------------------------------
TOTAL               9      0      4      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ######################
```

### Task 2

Create a Python function called `flatten_sequences` that receives two lists of integers. The function must return a new list containing:
- All elements from list 1
- Followed by all elements from list 2

#### Rules
- You must use `chain` from `itertools` module.
- Do not use: `+` or `.extend()` or manual loops
- The function must return a new list
- If both inputs are empty, return an empty list
- You must import only the required functions. Do not import the entire module or use wildcard imports.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from itertools import chain


def flatten_sequences(x: list[int], y: list[int]) -> list[int]:
    if x == [] and y == []:
        return []
    return list(chain(x, y))
```

```python
### uncomment the line(s) below to debug

list1, list2 = [1, 2], [3, 4]
result = flatten_sequences(list1, list2)
print(f"The flattened sequence is {result}")
print("Execution finished", "\u2713")
```

Output:
```
The flattened sequence is [1, 2, 3, 4]
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch


def test_solution_2() -> None:
    try:
        import solution_2
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_2 import flatten_sequences
    except:
        raise ValueError("The module does not have the necessary function!")

    assert flatten_sequences([1, 2], [3, 4]) == [1, 2, 3, 4]
    assert flatten_sequences([], []) == []
    assert flatten_sequences([1], []) == [1]
    assert flatten_sequences([], [5]) == [5]

    # Ensure new list is returned (no aliasing)
    a = [1, 2]
    b = [3, 4]
    result = flatten_sequences(a, b)
    assert result is not a
    assert result is not b

    with patch("solution_2.chain") as mock_chain:
        mock_chain.return_value = [9, 8]

        result = flatten_sequences([1], [2])
        assert result == [9, 8]

        mock_chain.assert_called_once()
        assert mock_chain.called, \
            "Do not manually concatenate lists. Use `itertools.chain`."
```

```python
# validate the solution
solution_file = "solution_2.py"
test_file = "test_solution_2.py"
export_cell_content("solution_2", solution_file)
export_cell_content("tests_2", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################# coverage #############################################

[32m[32m[1m1 passed[0m[32m in 0.20s[0m[0m
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_2.py       5      0      2      0   100%
-----------------------------------------------------------
TOTAL               5      0      2      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0flake8 True

[0m
############################################## pylint #######################
```

### Task 3

Create a Python function called `gm_median` that receives a list containing elements of `int`, `float`, and `str` types.
The function must return a `tuple` (geometric mean, median) computed using the `statistics` module.

The function must:
- Ignore elements that are not numbers (int or float)
- Ignore numeric values that are less than or equal to 0
- If the original list is empty, return None
- If, after filtering, no valid numeric values remain, return None
- Otherwise, return a tuple (geometric_mean, median) computed using the statistics module

You must import only the required functions. Do not import the entire module or use wildcard imports.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from statistics import geometric_mean, median


def gm_median(x: list[int | float | str]) -> tuple[float, float] | None:
    ny: list[float | int] = []
    if not x:
        return None
    for nummer in x:
        if isinstance(nummer, (int, float)) and nummer > 0:
            ny.append(nummer)
    if not ny:
        return None
    return (geometric_mean(ny), median(ny))
```

```python
### uncomment the line(s) below to debug

numbers = [0.5, "ignore", 2.0]
stats = gm_median(numbers)
print(f'Geometric Mean: {stats[0]}')
print(f'Median: {stats[1]}')
print('Execution finished', u'\u2713')
```

Output:
```
Geometric Mean: 1.0
Median: 1.25
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch
import math

def test_solution_3() -> None:
    try:
        import solution_3
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_3 import gm_median
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [
        ([], None),                       # empty list
        ([0, -1, -2.5, "x"], None),       # all numeric <= 0
        (["x", 0, -3, "7"], None),        # strings ignored, <=0 ignored → nothing left

        ([2, "x", 0, 8, -3, "5"], (4.0, 5.0)),      # valid numbers [2, 8]
        (["x", 7, 0, -1], (7.0, 7.0)),              # single valid number
        ([0.5, "ignore", 2.0], (1.0, 1.25)),        # float handling
    ]

    for _in, _out in test_cases:
        _res = gm_median(_in)
        
        if _out == None:
            assert _res is None, f"The function with input `{_in}` should return None, \
        but returned `{_res}`."
            continue

        assert _res is not None, f"The function with input `{_in}` returned None unexpectedly."

        cond1 = math.isclose(_res[0], _out[0], abs_tol=1e-1)
        cond2 = math.isclose(_res[1], _out[1], abs_tol=1e-1)

        assert cond1 and cond2, f"The function with input `{_in}` should return the value \
    `{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."

    with patch("solution_3.geometric_mean") as mock_gm, \
        patch("solution_3.median") as mock_med:

        mock_gm.return_value = 123.0
        mock_med.return_value = 456.0
        assert gm_median([1, 2, 3]) == (123.0, 456.0)
        mock_gm.assert_called_once()
        mock_med.assert_called_once()
        assert (
            mock_gm.called and mock_med.called
        ), "Do not use `import statistics`. Import the specific functions from the module."
```

```python
# validate the solution
solution_file = "solution_3.py"
test_file = "test_solution_3.py"
export_cell_content("solution_3", solution_file)
export_cell_content("tests_3", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.28s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_3.py      11      0      8      0   100%
-----------------------------------------------------------
TOTAL              11      0      8      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ######################
```

### Task 4

Create a Python function called `random_mobile_number_generator` that randomly generates a mobile number (`str`).

The number must:
- Start with one of: `"0"`, `"46"`, `"+46"`
- The remaining part must contain exactly 9 digits
- The first digit of the remaining part must be `"3"` or `"7"`
- The remaining 8 digits must be between `0` and `9`

#### Additional requirement

The function must accept an optional parameter: `seed: int | None = None`
- If `seed` is provided, the output must be deterministic.
- Use `random.Random(seed)` to achieve this.

You must import and use module `random` in your solution. You must import only the required functions. Do not import the entire module or use wildcard imports.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from random import Random


def random_mobile_number_generator(seed: int | None = None) -> str:
    rng = Random(seed)

    prefix = rng.choice(["0", "46", "+46"])
    first_digit = rng.choice(["3", "7"])
    remaining_digits = "".join(str(rng.randint(0, 9)) for _ in range(8))

    return prefix + first_digit + remaining_digits
```

```python
### uncomment the line(s) below to debug

print(random_mobile_number_generator())
print(random_mobile_number_generator(seed=123))
print(random_mobile_number_generator(seed=123))
print('Execution finished', u'\u2713')
```

Output:
```
46329458285
0716410688
0716410688
Execution finished ✓
```

```python
# test cell
# testing the probability
import collections
import random
from unittest.mock import patch


def test_solution_4() -> None:
    try:
        import solution_4
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_4 import random_mobile_number_generator
    except:
        raise ValueError("The module does not have the necessary function!")

    n1 = random_mobile_number_generator(seed=123)
    n2 = random_mobile_number_generator(seed=123)
    n3 = random_mobile_number_generator(seed=124)

    assert n1 == n2, "When seed is provided, output must be deterministic (same seed => same output)."
    assert n1 != n3, "Different seeds should (very likely) yield different outputs."

    idx_digit = random.randint(1, 8)
    number_samples = 100000
    count_prefix = collections.defaultdict(int)
    count_digit = collections.defaultdict(int)
    valid_prefixes = ["0", "46", "+46"]

    for i in range(number_samples):
        number = random_mobile_number_generator()
        res = [number.startswith(prefix) for prefix in valid_prefixes]
        assert any(res) == True, f"Generated number {number} is invalid due to incorrect \
    prefix. Check the rules again."
        prefix = valid_prefixes[res.index(True)]
        count_prefix[prefix] += 1
        remaining_part = number.replace(prefix, '', 1)
        assert remaining_part.isdigit() == True, f"Generated number {number} is invalid \
    due to incorrect digits. Check the rules again."
        assert len(remaining_part) == 9, f"Generated number {number} is invalid due to \
    incorrect length. Check the rules again."
        assert remaining_part[0] in ['3', '7'], f"Generated number {number} is invalid \
    due to `{remaining_part[0]}` being first digit after prefix. Check the rules again."
        count_digit[remaining_part[idx_digit]] += 1

    for _, count in count_prefix.items():
        assert abs(count / number_samples - 1/3) < 0.01

    for _, count in count_digit.items():
        assert abs(count / number_samples - 1/10) < 0.01
```

```python
# validate the solution
solution_file = "solution_4.py"
test_file = "test_solution_4.py"
export_cell_content("solution_4", solution_file)
export_cell_content("tests_4", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 8.75s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_4.py       7      0      0      0   100%
-----------------------------------------------------------
TOTAL               7      0      0      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0flake8 True

[0m
############################################## pylint #######################
```

### Task 5

Create a Python function called `lottery_numbers` that receives three parameters:
1. number: how many numbers should be generated
2. upper: the upper bound (inclusive)
3. lower: the lower bound (inclusive)

The function must generate a list of unique integer numbers with exactly number elements, each between lower and upper (inclusive).

#### Validation rules
Return None if any of the following is true:
- lower < 0
- upper < 1
- number < 1
- (upper - lower) < (number - 1) (not enough unique integers available)

You must import and use module `random`. You must import only the required functions.
Do not import random and do not use wildcard imports.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

```python
# solution cell
from random import Random


def lottery_numbers(number: int, upper: int, lower: int) -> list[int] | None:
    if lower < 0 or upper < 1 or number < 1 or (upper - lower) < (number - 1):
        return None
    rng = Random()
    lista: list[int] = []
    nytt = 0
    while len(lista) < number:
        nytt = rng.randint(lower, upper)
        if nytt not in lista:
            lista.append(nytt)
    return lista
```

```python
### uncomment the line(s) below to debug

result = lottery_numbers(5, 50, 10)
print(f'Generating 5 numbers between 10 and 50 is {result}')
print('Execution finished', u'\u2713')
```

Output:
```
Generating 5 numbers between 10 and 50 is [38, 31, 14, 22, 13]
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch


def test_solution_5() -> None:
    try:
        import solution_5
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_5 import lottery_numbers
    except:
        raise ValueError("The module does not have the necessary function!")

    invalid_inputs = [
        (5, 50, -1),     # lower < 0
        (5, 0, 0),       # upper < 1
        (0, 10, 0),      # number < 1
        (5, 3, 0),       # not enough unique values
        (30, -1, 0),     # upper < 1 (also negative)
    ]
    for _in in invalid_inputs:
        _res = lottery_numbers(*_in)
        assert _res is None, \
            f"The function received the input `{_in}` and should return None, " \
            f"but returned `{_res}`."

    valid_inputs = [
        (5, 50, 10),
        (30, 100, 0),
        (2, 1, 0),       # edge case: population size == number
    ]
    for _in in valid_inputs:
        _res = lottery_numbers(*_in)
        assert _res is not None, f"The function returned None for valid input `{_in}`."
        assert len(_res) == _in[0], f"Wrong length for input `{_in}`. Output: `{_res}`."
        assert len(_res) == len(set(_res)), \
            f"Repeated elements for input `{_in}`. Output: `{_res}`."
        assert min(_res) >= _in[2], \
            f"Element below lower bound for input `{_in}`. Output: `{_res}`."
        assert max(_res) <= _in[1], \
            f"Element above upper bound for input `{_in}`. Output: `{_res}`."
```

```python
# validate the solution
solution_file = "solution_5.py"
test_file = "test_solution_5.py"
export_cell_content("solution_5", solution_file)
export_cell_content("tests_5", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################# coverage #############################################

[32m[32m[1m1 passed[0m[32m in 0.33s[0m[0m
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_5.py      12      0      6      0   100%
-----------------------------------------------------------
TOTAL              12      0      6      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0flake8 True

[0m
############################################## pylint #######################
```

### Task 6

Create a Python function called `datetime_validator` that receives a string representing a date-time and converts it into a `datetime` object.

Accepted formats (exact):
1. "January 15, 2022, 12:30:45" (full month name)
2. "Jan 15, 2022, 12:30:45" (abbreviated month name)

Rules:
- Ignore leading/trailing whitespace in the input (e.g., "  January 15, 2022, 12:30:45  " is valid).
- If parsing succeeds in either format, return the `datetime` object.
- If parsing fails, return `None`.

Implementation requirements:
- You must use [`datetime.strptime()`](https://docs.python.org/3/library/datetime.html#datetime.datetime.strptime) and a `try/except` block. Check the [format codes](https://docs.python.org/3/library/datetime.html#strftime-and-strptime-format-codes).
- Import requirement (direct style): `from datetime import datetime`
- Do not use import `datetime` and do not use wildcard imports.

```python
# solution cell
from datetime import datetime


def datetime_validator(s: str) -> datetime | None:
    s = s.strip()

    formats = [
        "%B %d, %Y, %H:%M:%S",
        "%b %d, %Y, %H:%M:%S",
    ]

    for fmt in formats:
        try:
            return datetime.strptime(s, fmt)
        except ValueError:
            continue

    return None
```

```python
# uncomment the line(s) below to debug

date_input = "January 15, 2022, 12:30:45"
print(f'{date_input} is equivalent to {datetime_validator(date_input)} datetime object')
date_input = "Jan 15, 2022, 12:30:45"
print(f'{date_input} is equivalent to {datetime_validator(date_input)} datetime object')
print('Execution finished', u'\u2713')
```

Output:
```
January 15, 2022, 12:30:45 is equivalent to 2022-01-15 12:30:45 datetime object
Jan 15, 2022, 12:30:45 is equivalent to 2022-01-15 12:30:45 datetime object
Execution finished ✓
```

```python
# test cell
from datetime import datetime as dt
from unittest.mock import patch


def test_solution_6() -> None:
    try:
        import solution_6
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_6 import datetime_validator
    except:
        raise ValueError("The module does not have the necessary function!")

    valid_cases = [
        ("January 15, 2022, 12:30:45", dt(2022, 1, 15, 12, 30, 45)),
        ("Jan 15, 2022, 12:30:45", dt(2022, 1, 15, 12, 30, 45)),
        ("  January 15, 2022, 12:30:45  ", dt(2022, 1, 15, 12, 30, 45)),
        ("Feb 29, 2024, 00:00:00", dt(2024, 2, 29, 0, 0, 0)),  # leap year check
    ]

    for _in, _out in valid_cases:
        _res = datetime_validator(_in)
        assert _res == _out, \
            f"Input `{_in}` should return `{_out}`, but returned `{_res}`."

    invalid_cases = [
        "2021-01-01 12:30:45",          # wrong format entirely
        "September 31, 2021, 14:45:15", # invalid date
        "Jan 15 2022 12:30:45",         # missing commas
        "",                             # empty string
    ]

    for bad in invalid_cases:
        _res = datetime_validator(bad)
        assert _res is None, \
            f"Input `{bad}` should return None, but returned `{_res}`."

    with patch("solution_6.datetime") as mock_datetime:
        mock_datetime.strptime.return_value = dt(2000, 1, 1, 0, 0, 0)
        result = datetime_validator("January 15, 2022, 12:30:45")
        assert result == dt(2000, 1, 1, 0, 0, 0)
        mock_datetime.strptime.assert_called()
        assert mock_datetime.strptime.called, \
            "You must use `datetime.strptime()` to parse the string."
```

```python
# validate the solution
solution_file = "solution_6.py"
test_file = "test_solution_6.py"
export_cell_content("solution_6", solution_file)
export_cell_content("tests_6", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.27s[0m[0m

[0m
############################################# coverage #############################################
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_6.py      10      0      2      0   100%
-----------------------------------------------------------
TOTAL              10      0      2      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0flake8 True

[0m
############################################## pylint #######################
```

### Task 7

Create a Python function called `most_common_words` that receives:
1. words: list of strings
2. n: integer

The function must return a list of the n most common words in descending frequency order. Each element in the returned list must be a tuple: (word, count)

#### Rules
- If n < 1, return `None`
- If words is empty, return `None`
- If n is greater than the number of unique words, return all unique words
- Words are case-sensitive
- You must use `Counter` from `collections` module. Do not manually implement counting logic

```python
# solution cell
from collections import Counter


def most_common_words(words: list[str], n: int) -> list[tuple[str, int]] | None:
    if n < 1 or words == []:
        return None
    c = Counter(words)
    return c.most_common(n)
```

```python
# uncomment the line(s) below to debug

n = 2
seq_input = ["apple", "banana", "apple", "orange", "banana", "apple"]
out = most_common_words(seq_input, n)
print(f'{n} most common words: ', out)
print('Execution finished', u'\u2713')
```

Output:
```
2 most common words:  [('apple', 3), ('banana', 2)]
Execution finished ✓
```

```python
# test cell
from unittest.mock import patch


def test_solution_7() -> None:
    try:
        import solution_7
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_7 import most_common_words
    except:
        raise ValueError("The module does not have the necessary function!")

    assert most_common_words([], 3) is None
    assert most_common_words(["a", "b"], 0) is None

    words = ["apple", "banana", "apple", "orange", "banana", "apple"]

    result = most_common_words(words, 2)
    expected = [("apple", 3), ("banana", 2)]

    assert result == expected, \
        f"Expected {expected}, but got {result}."

    # n larger than unique words
    result = most_common_words(words, 10)
    assert len(result) == 3

    with patch("solution_7.Counter") as mock_counter:
        mock_instance = mock_counter.return_value
        mock_instance.most_common.return_value = [("x", 5)]

        result = most_common_words(["x", "x"], 1)
        assert result == [("x", 5)]

        mock_counter.assert_called_once()
        mock_instance.most_common.assert_called_once()

        assert mock_counter.called, \
        "Do not use `import collections`. Import the specific functions from the module."
```

```python
# validate the solution
solution_file = "solution_7.py"
test_file = "test_solution_7.py"
export_cell_content("solution_7", solution_file)
export_cell_content("tests_7", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" coverage ".center(100, "#"))
!coverage report -m {solution_file}
if _exit_code != 0:
    passed = False
print(f"coverage {passed}")

print('\n\033[0m')
print(" flake8 ".center(100, "#"))
!flake8 {solution_file}
if _exit_code != 0:
    passed = False
print(f"flake8 {passed}")

print('\n\033[0m')
print(" pylint ".center(100, "#"))
!pylint {solution_file}
if _exit_code != 0:
    passed = False
print(f"pylint {passed}")

print('\n\033[0m')
print(" mypy ".center(100, "#"))
!mypy {solution_file}
if _exit_code != 0:
    passed = False
print(f"mypy {passed}")

print('\n\033[0m')
print(" isort ".center(100, "#"))
!isort --check --diff {solution_file}
if _exit_code != 0:
    passed = False
print(f"isort {passed}")

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)
Latest saved date of the file: 2026-04-29 18:52:03.800149 (5 days ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################# coverage #############################################

[32m[32m[1m1 passed[0m[32m in 0.20s[0m[0m
Name            Stmts   Miss Branch BrPart  Cover   Missing
-----------------------------------------------------------
solution_7.py       6      0      2      0   100%
-----------------------------------------------------------
TOTAL               6      0      2      0   100%
coverage True

[0m
############################################## flake8 ##############################################
0
flake8 True

[0m
############################################## pylint ######################
```