# EEN065 - Applied object-oriented programming


Teacher: [Carlos Natalino](https://www.chalmers.se/en/persons/carda/) / Examiner: [Paolo Monti](https://www.chalmers.se/en/persons/mpaolo/)

[Canvas course page](https://chalmers.instructure.com/courses/39015)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before submitting this assignment, please ensure that everything runs correctly.

First, **restart the kernel**, then **run all cells**.
Next, verify that all tests execute successfully.

Note that if any problem produces an error, the subsequent problems **will not** be tested.

In case of discrepancies between the problem description and the tests, you should resolve the issue with the tests in mind, as they determine correctness.

There are four types of cells in this notebook:
1.	*Solution cells*: These are the cells where you write your answers or modify the existing code to solve the problem.
2.	*Debug cells*: These cells are for debugging your solution. You may add or edit code here, but ensure that it does not introduce any errors.
3. *Test cells:* These cells run **automated tests** to verify whether your solution is **functionally correct**.

- If all tests pass, you will see an output similar to:
  - `1 passed`

- If a test fails, you will see an **error message** explaining:
  - Which input was tested
  - What output was expected
  - What your function actually returned

You will see following output:

- **tests:** Runs the test cells with coverage enabled.  
  - `1 passed` indicates success.

At the end:
- **`tests passed ✓`** → Your solution is correct *and* meets all quality requirements.
- **`your code presented some issue`** → Review the messages above, fix the reported issues, and rerun the cell.

**You must pass both functional tests and code quality checks to receive full credit.**

Common errors and what they mean:
- **`You did not execute your solution cell!`**  
  → You forgot to run the solution cell before running the test cell.
- **`The module does not have the necessary function!`**  
  → The required function is missing or named incorrectly.
- **AssertionError**  
  → Your function ran, but it returned an incorrect result for at least one test case.

**Fix all test errors before moving on**, as later problems will not be tested if an earlier one fails.

**Delete** the line `raise NotImplementedError()` from any problem you solve.

**Do not add or remove any cells in this file.** All required cells are already provided.

To execute a cell, select it and press **CTRL+Enter** (Windows), **CMD+Enter** (macOS), or click the **Run cell** button.

## Assignment Week 2

**Observations:**
For the problems in this assignment, you should:
- use only the math operations studied in the lecture notes (i.e., no advanced functions allowed).
- not use external modules, i.e., `import` command is not allowed in your solutions (with exception of the module `typing`).

**You should need only Chapters 1-3 of the lecture notes to solve all the tasks in this list.**

**Preparation:** Run the cell below every time you start working on this file, and every time you restart the kernel.

```python
%load_ext autoreload
import os

if not os.path.exists("utils.py"):
    raise ValueError("It seems this file is in the wrong folder. "
          "Make sure to place it in the `programming-assignments` folder/project.")

# check if this file is within the same folder as the `utils.py` file.
try:
    from utils import export_cell_content
except:
    import sys
    raise ValueError("Did you select the right kernel that starts with venv?")

filename = "assignment_week_02_public.ipynb"
assert filename != None, "The file name must be set."
if not os.path.exists(filename):
    raise ValueError(f"The value of the `filename` variable is wrong. \
The file `{filename}` does not exist. Write the correct name of this file!")

os.environ["NOTEBOOK_NAME"] = os.path.join(os.getcwd(), filename)
```

### Task 1

Create a Python function called `compute_expression` that receives a two positive integers $a$ and $b$ and returns the following expression:

\begin{equation*}
a + \left\{ \dfrac{b}{a+ \left \{ \dfrac{b}{a+b} \right \} } \right\}. 
\end{equation*}


```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def compute_expression(a: int, b: int) -> float: 
    result = a + (b / (a + ( b / (a + b))))
    return result
```

```python
# uncomment the line(s) below to debug

a = 7
b = 2
c = compute_expression(a, b)
print(f"The value of the expression given a = {a} and b = {b} is {c}")
print("Execution finished", "\u2713")
```

Output:
```
The value of the expression given a = 7 and b = 2 is 7.276923076923077
Execution finished ✓
```

```python
# test cell
def test_solution_1() -> None:
    try:
        import solution_1
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_1 import compute_expression
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((2, 4), 3.5), ((3,9), 5.4), ((7, 2), 7.276923076923077)]

    for _in, _out in test_cases:
        _res = compute_expression(*_in)
        assert _res == _out, f"The function with input `{_in}` should return the value \
    `{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_1.py"
test_file = "test_solution_1.py"
export_cell_content("solution_1", solution_file)
export_cell_content("tests_1", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 11:43:17.344679 (7 seconds ago)
Latest saved date of the file: 2026-03-31 11:43:17.344679 (8 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.09s[0m[0m

[0m
############################################## result ##############################################
[m tests passed ✓ [0m
```

### Task 2

Create a Python function called `print_computed_expression` that receives a two positive integers $a$ and $b$ and prints a $\underline{\textbf{string}}$ representing the expression:

\begin{equation*}
a + \left\{ \dfrac{b}{a+ \left \{ \dfrac{b}{a+b} \right \} } \right\}. 
\end{equation*}

Note that the output is returned as a string rather than an integer or float.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
# solution cell
def print_computed_expression(a: int, b: int) -> str:
    result = str(a + (b / (a + ( b / (a + b)))))
    print(result)
    return result
```

```python
# uncomment the line(s) below to debug

a = 7
b = 2
print(f"The computed expression is ")
print_computed_expression(a, b)
print("Execution finished", "\u2713")
```

Output:
```
The computed expression is 
7.276923076923077
Execution finished ✓
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_2() -> None:
    try:
        import solution_2
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_2 import print_computed_expression
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((2, 4), "3.5"), ((3,9), "5.4"), ((7, 2), "7.276923076923077")]

    for _in, _out in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = print_computed_expression(*_in)
            assert fakeOutput.getvalue().strip() == _out, f"The function should print \
    the value `{_out}` but printed the value `{fakeOutput.getvalue().strip()}`."
```

```python
# validate the solution
solution_file = "solution_2.py"
test_file = "test_solution_2.py"
export_cell_content("solution_2", solution_file)
export_cell_content("tests_2", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

### Task 3

Create a function called `cylinder_volume` that, based on two parameters representing the radius of the base and height of the cylinder (integer or float), prints its volume upto four decimal places.

Remember that the volume $V$ of a cylinder given its base radius $r$ and height $h$ is given by:
\begin{equation*}
V = \pi r^2 h. 
\end{equation*}

**Observations:**
- You should assume $\pi=3.14$.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-3, points: 0, grade: false, solution: true -->
```python
# solution cell
def cylinder_volume(r: int | float, h: int | float) -> float:
    pi = 3.14
    v = pi * (r**2) * h
    print(f"{v:.4f}")
```

```python
# uncomment the line(s) below to debug

r, h = 0.11, 0.3
print(f'The volume of a cylinder with radius {r} and height {h} is')
cylinder_volume(r, h)
print("Execution finished", "\u2713")
```

Output:
```
The volume of a cylinder with radius 0.11 and height 0.3 is
0.0114
Execution finished ✓
```

<!-- grade_id: cell-test-1-3, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_3() -> None:
    try:
        import solution_3
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_3 import cylinder_volume
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((2.1, 4.3), "59.5438"), ((10, 20), "6280.0000"), ((0.11, 0.3), "0.0114")]

    for _in, _out in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = cylinder_volume(*_in)
            assert fakeOutput.getvalue().strip() == _out, f"The function should print \
    the value `{_out}` but printed the value `{fakeOutput.getvalue().strip()}`."
```

```python
# validate the solution
solution_file = "solution_3.py"
test_file = "test_solution_3.py"
export_cell_content("solution_3", solution_file)
export_cell_content("tests_3", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 12:13:30.220186 (1 minutes ago)
Latest saved date of the file: 2026-03-31 12:13:30.220186 (1 minutes ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.24s[0m[0m

[0m
############################################## result ##############################################
[m tests passed ✓ [0m
```

### Task 4

Create a function named `greetings` that receives:
- the salutation of a person (string),
- the name of a person (string),
- the surname of a person (string),
- the age of a person (integer)
- the location of a person (string),
and returns the following greeting message.

For instance, for a person of name 'John', surname 'Doe', salutation 'Mr.', age 39 and location 'paris', the function should return the following message which contains two lines:

`Hello from Paris!`

`I'm Mr. John Doe, age 39.`

<!-- Greetings! I am Mr. John Doe, and I am 25 years old. -->

The name, surname, salutation and location should always be capitalized. For instance, for a person of name 'jAne', surname 'DoE', salutation 'mS.', age 32 and location 'bErlin', the function should return a message similar to the following one:

`Hello from Berlin!`

`I'm Ms. Jane Doe, age 32.`

**Observations:**
- You should use `f-string` and `capitalize()` to solve this problem.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-4, points: 0, grade: false, solution: true -->
```python
# solution cell
def greetings(salutation, name, surname, age, location) -> str:
    salutation = salutation.capitalize()
    name = name.capitalize()
    surname = surname.capitalize()
    location = location.capitalize()
    line1 = f"Hello from {location}!"
    line2 = f"I'm {salutation} {name} {surname}, age {age}."
    return line1 + "\n" + line2
```

```python
# uncomment the line(s) below to debug

g = greetings('ms.', 'jAne', 'DoE', 32, 'bErlin')
print(g)
print("Execution finished", "\u2713")
```

Output:
```
Hello from Berlin!
I'm Ms. Jane Doe, age 32.
Execution finished ✓
```

<!-- grade_id: cell-test-1-4, points: 1, grade: true, solution: false -->
```python
# test cell
def test_solution_4() -> None:
    try:
        import solution_4
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_4 import greetings
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(('MR.', 'joHn', 'DOe', 39, 'pARiS'), 'Hello from Paris!\nI\'m Mr. John Doe, age 39.'), \
    (('ms.', 'jAne', 'DoE', 32, 'bErLiN'), 'Hello from Berlin!\nI\'m Ms. Jane Doe, age 32.')]

    for _in, _out in test_cases:
        _res = greetings(*_in)
        assert _res == _out, f"The function with input `{_in}` should return the value \
    `{_out}` of type `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_4.py"
test_file = "test_solution_4.py"
export_cell_content("solution_4", solution_file)
export_cell_content("tests_4", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 12:48:28.067688 (8 seconds ago)
Latest saved date of the file: 2026-03-31 12:48:28.067688 (8 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################## result ##############################################
[m tests passed ✓ [0m

[32m[32m[1m1 passed[0m[32m in 0.09s[0m[0m
```

### Task 5

Create a Python function called `print_return_product` that receives two numbers (integer or float).
The function should:
1. print a string that represents the product of the inputs; and
2. return the product of the inputs.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def print_return_product(x: int | float, y: int | float) -> str:
    result = x * y
    print(result)
    return result
```

```python
# uncomment the line(s) below to debug

x, y = 4.2, 2.3
print(f'the product of {x} and {y} is:')
r = print_return_product(x, y)
print(f'or {r}')
print("Execution finished", "\u2713")
```

Output:
```
the product of 4.2 and 2.3 is:
9.66
or 9.66
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_5() -> None:
    try:
        import solution_5
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_5 import print_return_product
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((8.6352456, 52.18234), 450.607321882704, "450.607321882704"), \
        ((4.2, 2.3), 9.66, "9.66")]

    for _in, _out, _print in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = print_return_product(*_in)
            assert _res == _out, f"The function should return the value `{_out}` of type \
        `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
            assert fakeOutput.getvalue().strip() == _print, f"The function should print \
    the value `{_print}` but printed the value `{fakeOutput.getvalue().strip()}`."
```

```python
# validate the solution
solution_file = "solution_5.py"
test_file = "test_solution_5.py"
export_cell_content("solution_5", solution_file)
export_cell_content("tests_5", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 14:08:37.938020 (12 seconds ago)
Latest saved date of the file: 2026-03-31 14:08:37.938020 (12 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[0m
############################################## result ##############################################
[m tests passed ✓ [0m

[32m[32m[1m1 passed[0m[32m in 0.21s[0m[0m
```

### Task 6

Create a Python function called `return_product` that receives two numbers (integer or float).
The function should return a $\underline{\textbf{string}}$ that represents the product of the inputs. Note that the output is returned as a string rather than an integer or float.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
# solution cell
def return_product(x: int | float, y: int | float) -> str:
    result = x * y
    print(result)
    return str(result)
```

```python
# uncomment the line(s) below to debug

x, y = 4.2, 2.3
r = return_product(x, y)
print(f'the product of {x} and {y} is: {r}')
print("Execution finished", "\u2713")
```

Output:
```
9.66
the product of 4.2 and 2.3 is: 9.66
Execution finished ✓
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_6() -> None:
    try:
        import solution_6
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_6 import return_product
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [((8.6352456, 52.18234), "450.607321882704"), \
        ((4.2, 2.3), "9.66")]

    for _in, _out in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = return_product(*_in)
            assert _res == _out, f"The function should return the value `{_out}` of type \
        `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_6.py"
test_file = "test_solution_6.py"
export_cell_content("solution_6", solution_file)
export_cell_content("tests_6", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 14:11:47.507970 (11 seconds ago)
Latest saved date of the file: 2026-03-31 14:11:47.507970 (11 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.10s[0m[0m

[0m
############################################## result ##############################################
[m tests passed ✓ [0m
```

### Task 7

Create a function named `coin_change` that receives an amount (integer or float) and returns the equivalent amount as a tuple of (hundreds, fifties, tens, fives, ones) (in this order). 

You need to round down the outputs to nearest integer.

For instance:

- Input: `179` / Output: `(1, 1, 2, 1, 4)`
- Input: `179.5` / Output: `(1, 1, 2, 1, 4)`
- Input: `365` / Output: `(3, 1, 1, 1, 0)`

**Observations:**
- Remember that you cannot use `import`. Use only the math operations in Chapter 3.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-2, points: 0, grade: false, solution: true -->
```python
# solution cell
def coin_change(amount: int | float) -> tuple[int, int, int, int, int]:
    remaining = int(amount)
    hundreds = remaining // 100
    remaining = remaining % 100
    fifties = remaining // 50
    remaining = remaining % 50
    tens = remaining // 10
    remaining = remaining % 10
    fives = remaining // 5
    remaining = remaining % 5
    ones = remaining
    return (hundreds, fifties, tens, fives, ones)
```

```python
# uncomment the line(s) below to debug

amount = 179.5
v = coin_change(amount)
print(f"{amount} amount is equivalent to {v[0]} hundreds, {v[1]} fifties, {v[2]} tens, {v[3]} fives, {v[4]} ones")
print("Execution finished", "\u2713")
```

Output:
```
179.5 amount is equivalent to 1 hundreds, 1 fifties, 2 tens, 1 fives, 4 ones
Execution finished ✓
```

<!-- grade_id: cell-test-2-2, points: 1, grade: true, solution: false -->
```python
# test cell
from unittest.mock import patch
from io import StringIO

def test_solution_7() -> None:
    try:
        import solution_7
    except:
        raise ValueError("You did not execute your solution cell!")
    try:
        from solution_7 import coin_change
    except:
        raise ValueError("The module does not have the necessary function!")

    test_cases = [(179, (1, 1, 2, 1, 4)), (179.5, (1, 1, 2, 1, 4)), (365.25, (3, 1, 1, 1, 0))]

    for _in, _out in test_cases:
        with patch("sys.stdout", new=StringIO()) as fakeOutput:
            _res = coin_change(_in)
            assert _res == _out, f"The function should return the value `{_out}` of type \
        `{type(_out)}`\n but returned the value `{_res}` of type `{type(_res)}`."
```

```python
# validate the solution
solution_file = "solution_7.py"
test_file = "test_solution_7.py"
export_cell_content("solution_7", solution_file)
export_cell_content("tests_7", test_file)
passed = True

print('\n\033[0m')
print(" tests ".center(100, "#"))

!coverage run -m pytest -q {test_file}
if _exit_code != 0:
    passed = False

print('\n\033[0m')
print(" result ".center(100, "#"))
if passed:
    print('\033[m tests passed', u'\u2713', "\033[0m")
else:
    raise AssertionError('\033[91m your code presented some issue, check messages '\
        'above', u'\u274C', "\033[0m")
```

Output:
```
Latest saved date of the file: 2026-03-31 14:21:19.752735 (8 seconds ago)
Latest saved date of the file: 2026-03-31 14:21:19.752735 (8 seconds ago)

[0m
############################################## tests ###############################################
[32m.[0m[32m                                                                        [100%][0m
[32m[32m[1m1 passed[0m[32m in 0.08s[0m[0m

[0m
############################################## result ##############################################
[m tests passed ✓ [0m
```